---
title: '2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: f0aae17bcc   
```


new chapter


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


